# Tests for Activity 10

## Tests: Chinook

<!-- grade_id: 1d43c2, points: 4, grade: true, solution: false -->
```python
# Test for part 1
from nose.tools import assert_equal, assert_in
from csc270tools import get_table, get_script

# Running the code in act10Chinook.sql to get the table and your code for this part
table, header  = get_table( 'act10Chinook.sql', 'Chinook.db', 0, header=True)
script = get_script('act10Chinook.sql', 0).upper()

# Checking that a Window Function is used
assert_in('PARTITION BY',script, 'checking for a window function')

# Partially checking your solution
assert_equal(header, ['Album Title', 'Track Name', 'Track Length', 'Album Average Track Length'], 'Checking the column names')
assert_equal(('In Your Honor [Disc 1]', 'Hell', 117080, 240929.5), table[0], 'Checking the first row')
assert_equal(('The Colour And The Shape', 'New Way Home', 342230, 216560.76923076922), table[-1], 'Checking the last row')
assert_equal(len(table), 44, 'Checking the number of rows')

# There are hidden tests that will completely check your solution
```

<!-- grade_id: 4131eb, points: 4, grade: true, solution: false -->
```python
# Test for part 2
from nose.tools import assert_equal, assert_in, assert_not_in
from csc270tools import get_table, get_script

# Running the code in act10Chinook.sql to get the table and your code for this part
table, header  = get_table( 'act10Chinook.sql', 'Chinook.db', 1, header=True)
script = get_script('act10Chinook.sql', 1).upper()

# Checking that a CTE is used instead of a Window Function
assert_not_in('PARTITION BY', script, 'checking for a window function')
assert_in('WITH', script, 'checking for a CTE')

# Partially checking your solution
assert_equal(header, ['Album Title', 'Track Name', 'Track Length', 'Album Average Track Length'], 'Checking the column names')
assert_equal(('In Your Honor [Disc 1]', 'Hell', 117080, 240929.5), table[0], 'Checking the first row')
assert_equal(('The Colour And The Shape', 'New Way Home', 342230, 216560.76923076922), table[-1], 'Checking the last row')
assert_equal(len(table), 44, 'Checking the number of rows')

# There are hidden tests that will completely check your solution
```

<!-- grade_id: d9f7ed, points: 2, grade: true, solution: false -->
```python
# Test for part 3
from nose.tools import assert_equal
from csc270tools import get_table

# Running the code in act10Chinook.sql to get the table and your code for this part
table, header  = get_table( 'act10Chinook.sql', 'Chinook.db', 2, header=True)

# Partially checking your solution
assert_equal(header, ['Album Title', 'Track Name', 'Track Length', 'Album Average Track Length', 'Album Min Track Length', 'Album Max Track Length'], 'Checking the column names')
assert_equal(('In Your Honor [Disc 1]', 'Hell', 117080, 240929.5, 117080, 352078), table[0], 'Checking the first row')
assert_equal(('The Colour And The Shape', 'New Way Home', 342230, 216560.76923076922, 83487, 342230), table[-1], 'Checking the last row')
assert_equal(len(table), 44, 'Checking the number of rows')

# There are hidden tests that will completely check your solution
```

## Tests: USA.db

<!-- grade_id: 043f95, points: 4, grade: true, solution: false -->
```python
# Test for part 1
import sqlite3
from nose.tools import assert_equal, assert_in

# Setting up sqlite3 for unit tests
con = sqlite3.connect('USA.db')
cur = con.cursor()

# Checking that the new view exists with the correct name
views = [t[0] for t in cur.execute("SELECT tbl_name FROM sqlite_master WHERE type = 'view';").fetchall()]
assert_in('President', views, 'Checking for the new view')

# Checking the column names
view = cur.execute("SELECT * FROM President;").fetchall()
columns = [c[0] for c in cur.description]
assert_equal(columns, ['politician_id', 'firstname', 'middlename', 'lastname', 'start_date', 'end_date'], 'Checking the column names')

# Checking some rows of the View
assert_equal((1, 'George', None, 'Washington', '1789-04-30', '1797-03-04'), view[0], 'Checking the first row')
assert_equal((6, 'John', 'Quincy', 'Adams', '1825-03-04', '1829-03-04'), view[5], 'Checking another row')

# There are hidden tests that will completely check your solution
```

<!-- grade_id: 9ff22f, points: 2, grade: true, solution: false -->
```python
# Test for part 2
import sqlite3
from nose.tools import assert_equal

# Setting up sqlite3 for unit tests
con = sqlite3.connect('USA.db')
cur = con.cursor()

# Checking for Obama
table = cur.execute("""
    SELECT firstname, middlename, lastname, suffix, birthdate, 
           city_of_birth, state_of_birth, country_of_birth
    FROM Politician
    WHERE lastname = 'Obama';
    """).fetchall()

view  = cur.execute("""
    SELECT firstname, middlename, lastname, start_date, end_date 
    FROM President
    WHERE lastname = 'Obama';
    """).fetchall()

assert_equal(table, 
             [('Barack', 'Hussein', 'Obama', 'II', '1961-08-04', 
               'Honolulu', 'Hawaii', 'United States')], 
             'Checking the Politician Table')
assert_equal(view, 
             [('Barack', 'Hussein', 'Obama', '2009-01-20', '2017-01-20')], 
             'Checking the President View')
```

<!-- grade_id: ea4e29, points: 3, grade: true, solution: false -->
```python
# Test for part 3
import sqlite3
from nose.tools import assert_equal, assert_in

# Setting up sqlite3 for unit tests
con = sqlite3.connect('USA.db')
cur = con.cursor()

# Checking that the new view exists with the correct name
views = [t[0] for t in cur.execute("SELECT tbl_name FROM sqlite_master WHERE type = 'view';").fetchall()]
assert_in('DiedInOffice', views, 'Checking for the new view')

# Checking the column names
view = cur.execute("SELECT * FROM DiedInOffice;").fetchall()
columns = [c[0] for c in cur.description]
assert_equal(columns, ['firstname', 'lastname', 'office', 'deathdate'], 'Checking the column names')

# Checking some rows of the View
assert_equal(('George', 'Clinton', 'Vice President', '1812-04-20'), view[0], 'Checking the first row')
assert_in(('Abraham', 'Lincoln', 'President', '1865-04-15'), view, 'Checking another row')

# There are hidden tests that will completely check your solution
```

<!-- grade_id: 29abbf, points: 5, grade: true, solution: false -->
```python
# Test for part 4
import sqlite3
from nose.tools import assert_equal, assert_in

# Setting up sqlite3 for unit tests
con = sqlite3.connect('USA.db')
cur = con.cursor()

# Checking that the new view exists with the correct name
views = [t[0] for t in cur.execute("SELECT tbl_name FROM sqlite_master WHERE type = 'view';").fetchall()]
assert_in('VicePresident', views, 'Checking for the new view')

# Checking the column names
view = cur.execute("SELECT * FROM VicePresident;").fetchall()
columns = [c[0] for c in cur.description]
assert_equal(columns, ['vice_president', 'president', 'start_date', 'end_date'], 'Checking the column names')

# Checking some rows of the View
assert_equal(('John Adams', 'George Washington', '1789-04-30', '1797-03-04'), view[0], 'Checking the first row')
assert_equal(('Thomas Jefferson', 'John Adams', '1797-03-04', '1801-03-04'), view[1], 'Checking the second row')
assert_equal(('Aaron Burr', 'Thomas Jefferson', '1801-03-04', '1805-03-04'), view[2], 'Checking the third row')
assert_equal(('George Clinton', 'Thomas Jefferson', '1805-03-04', '1809-03-04'), view[3], 'Checking the fourth row')
assert_equal(('George Clinton', 'James Madison', '1809-03-04', '1812-04-20'), view[4], 'Checking the fifth row')
assert_equal(('John Tyler', 'William Harrison', '1841-03-04', '1841-04-04'), view[11], 'Checking the twelfth row')
assert_in(('Andrew Johnson', 'Abraham Lincoln', '1865-03-04', '1865-04-15'), view, 'Checking another row')

# There are hidden tests that will completely check your solution
```