# Tests for Activity 8

## Tests: Chinook

<!-- grade_id: a09ff2, points: 1, grade: true, solution: false -->
```python
# Test for part 1
from nose.tools import assert_equal, assert_greater, assert_in, assert_not_in
from csc270tools import get_table, get_script

# Running the code in act8Chinook.sql to get the table and your code for this part
table  = get_table( 'act8Chinook.sql', 'Chinook.db', 0)
script = get_script('act8Chinook.sql', 0).upper()

# Checking that subqueries are used
assert_greater(script.count('SELECT'), 1, 'checking for subqueries')
assert_not_in(' JOIN ', script, 'checking for JOINs')

# Partially checking your solution
assert_in(('Physical Graffiti [Disc 1]',), table, 'Checking a row')

# There are hidden tests that will completely check your solution
```

<!-- grade_id: ae225c, points: 1, grade: true, solution: false -->
```python
# Test for part 2
from nose.tools import assert_equal, assert_greater, assert_not_in
from csc270tools import get_table, get_script

# Running the code in act8Chinook.sql to get the table and your code for this part
table  = get_table( 'act8Chinook.sql', 'Chinook.db', 1)
script = get_script('act8Chinook.sql', 1).upper()

# Checking that subqueries are used
assert_greater(script.count('SELECT'), 1, 'checking for subqueries')
assert_not_in(' JOIN ', script, 'checking for JOINs')

# Partially checking your solution
assert_equal(('edfrancis@yachoo.ca',), table[0], 'Checking the first row')

# There are hidden tests that will completely check your solution
```

<!-- grade_id: 93888a, points: 1, grade: true, solution: false -->
```python
# Test for part 3
from nose.tools import assert_equal, assert_greater, assert_in, assert_not_in
from csc270tools import get_table, get_script

# Running the code in act8Chinook.sql to get the table and your code for this part
table  = get_table( 'act8Chinook.sql', 'Chinook.db', 2)
script = get_script('act8Chinook.sql', 2).upper()

# Checking that subqueries are used
assert_greater(script.count('SELECT'), 1, 'checking for subqueries')
assert_not_in(' JOIN ', script, 'checking for JOINs')

# Partially checking your solution
assert_in(('Andrew', 'Adams'), table, 'Checking one row')
assert_equal(len(table), 3, 'Checking the number of rows')

# There are hidden tests that will completely check your solution
```

<!-- grade_id: 096530, points: 2, grade: true, solution: false -->
```python
# Test for part 4
from nose.tools import assert_equal, assert_greater, assert_in, assert_not_in
from csc270tools import get_table, get_script

# Running the code in act8Chinook.sql to get the table and your code for this part
table  = get_table( 'act8Chinook.sql', 'Chinook.db', 3)
script = get_script('act8Chinook.sql', 3).upper()

# Checking that subqueries are used
assert_greater(script.count('SELECT'), 1, 'checking for subqueries')
assert_not_in(' JOIN ', script, 'checking for JOINs')

# Partially checking your solution
assert_in(('Robert', 'King'), table, 'Checking one row')
assert_greater(len(table), 4, 'Checking the number of rows')

# There are hidden tests that will completely check your solution
```

<!-- grade_id: 12d45f, points: 1, grade: true, solution: false -->
```python
# Test for part 5
import sqlite3
from nose.tools import assert_equal, assert_in

# Setting up sqlite3 for unit tests
con = sqlite3.connect('Chinook.db')
cur = con.cursor()

# Checking the database has been appropriately modified
table = cur.execute("SELECT FirstName, LastName, Email, HireDate FROM Employee;").fetchall()
con.close()

assert_in(('Dave', 'Rosoff', 'dave@chinookcorp.com', '2024-02-29'), table, 'Checking the new row exists')
assert_equal(len(table), 9, 'Checking the number of rows in the Employee table')

```

<!-- grade_id: c8a34f, points: 2, grade: true, solution: false -->
```python
# Test for part 6
import sqlite3
from nose.tools import assert_equal, assert_in, assert_greater
from csc270tools import get_table, get_script

# Getting your code in act8Chinook.sql for this part
script = get_script('act8Chinook.sql', 4).upper()

# Checking that subqueries are used
assert_greater(script.count('SELECT'), 0, 'checking for subqueries')

# Setting up sqlite3 for unit tests
con = sqlite3.connect('Chinook.db')
cur = con.cursor()

# Checking the database has been appropriately modified
table = cur.execute("""
    SELECT e1.FirstName, e2.FirstName 
    FROM Employee AS e1 JOIN Employee AS e2
    ON e1.ReportsTo = e2.ReportsTo AND e1.Title = e2.Title;
    """).fetchall()
assert_in(('Dave', 'Nancy'), table, 'Checking that Rosoff and Nancy have same title and boss')

# There are additional hidden tests for this part.
```

<!-- grade_id: ffc541, points: 2, grade: true, solution: false -->
```python
# Test for part 7
import sqlite3
from nose.tools import assert_equal, assert_not_in, assert_greater
from csc270tools import get_table, get_script

# Getting your code in act8Chinook.sql for this part
script = get_script('act8Chinook.sql', 5).upper()

# Checking that subqueries are used
assert_greater(script.count('SELECT'), 0, 'checking for subqueries')

# Setting up sqlite3 for unit tests
con = sqlite3.connect('Chinook.db')
cur = con.cursor()

# Checking the database has been appropriately modified
rows = cur.execute("""
    SELECT b.Email 
    FROM Employee AS e JOIN Employee AS b
    ON e.ReportsTo = b.EmployeeId; 
    """).fetchall()

assert_not_in(('nancy@chinookcorp.com',), rows, 'Checking that Nancy have no reports')
assert_equal(rows.count(('dave@chinookcorp.com',)), 3, "Checking that number of Rosoff's reports")

# There are additional hidden tests for this part.
```

## Tests: More Harry Potter

<!-- grade_id: 28b12e, points: 1, grade: true, solution: false -->
```python
# Test for part 1
from nose.tools import assert_equal, assert_greater, assert_in, assert_not_in
from csc270tools import get_table, get_script

# Running the code in act8Potter.sql to get the table and your code for this part
table  = get_table( 'act8Potter.sql', 'Potter.db', 0)
script = get_script('act8Potter.sql', 0).upper()

# Checking that subqueries are used
assert_greater(script.count('SELECT'), 1, 'checking for subqueries')
assert_not_in(' JOIN ', script, 'checking for JOINs')

# Partially checking your solution
assert_in(('Harry Potter', 'James Potter'), table, 'Checking one row')
assert_equal(len(table), 28, 'Checking the number of rows')

# There are hidden tests that will completely check your solution
```

<!-- grade_id: 1a9f40, points: 2, grade: true, solution: false -->
```python
# Test for part 2
from nose.tools import assert_equal, assert_greater, assert_in, assert_not_in
from csc270tools import get_table, get_script

# Running the code in act8Potter.sql to get the table and your code for this part
table  = get_table( 'act8Potter.sql', 'Potter.db', 1)
script = get_script('act8Potter.sql', 1).upper()

# Checking that subqueries are used
assert_greater(script.count('SELECT'), 1, 'checking for subqueries')
assert_not_in(' JOIN ', script, 'checking for JOINs')

# Partially checking your solution
assert_in(('James Potter', 'Harry Potter', 'James Potter'), table, 'Checking the first row')
assert_in(('Septimus Weasley', 'Arthur Weasley', 'George Weasley'), table, 'Checking the last row')
assert_equal(len(table), 28, 'Checking the number of rows')

# There are hidden tests that will completely check your solution
```

<!-- grade_id: 95c645, points: 2, grade: true, solution: false -->
```python
# Test for part 3
from nose.tools import assert_equal, assert_greater, assert_in, assert_not_in
from csc270tools import get_table, get_script

# Running the code in act8Potter.sql to get the table and your code for this part
table  = get_table( 'act8Potter.sql', 'Potter.db', 2)
script = get_script('act8Potter.sql', 2).upper()

# Checking that subqueries are used
assert_greater(script.count('SELECT'), 1, 'checking for subqueries')
assert_not_in(' JOIN ', script, 'checking for JOINs')

# Partially checking your solution
assert_in(('Ginny',), table, 'Checking the first row')
assert_in(('George',), table, 'Checking the last row')
assert_equal(len(table), 5, 'Checking the number of rows')

# There are hidden tests that will completely check your solution
```