# Tests for Activity 6

## Tests: creating tables

<!-- grade_id: 2b1440, points: 1, grade: true, solution: false -->
```python
# Tests for creating the table Yotes
import sqlite3
from nose.tools import assert_equal, assert_in

# Setting up sqlite3 for unit tests
con = sqlite3.connect(':memory:')
cur = con.cursor()

# Running all the code in act6create.sql
with open('act6create.sql') as file:
    cur.executescript(file.read())

# Checking that the table exists 
tables = [t[0] for t in cur.execute("SELECT tbl_name FROM sqlite_master;").fetchall()]

assert_in('Yotes', tables, 'Checking the table has the correct name')

# Checking the column names
cur.execute("SELECT * FROM Yotes")
columns = {c[0] for c in cur.description}

assert_equal(columns, {'LastName', 'FirstName', 'Email', 'YoteId'})

con.close()
```

<!-- grade_id: b8780a, points: 1, grade: true, solution: false -->
```python
# Tests for creating the table Courses
import sqlite3
from nose.tools import assert_equal, assert_in

# Setting up sqlite3 for unit tests
con = sqlite3.connect(':memory:')
cur = con.cursor()

# Running all the code in act6create.sql
with open('act6create.sql') as file:
    cur.executescript(file.read())

# Checking that the table exists 
tables = [t[0] for t in cur.execute("SELECT tbl_name FROM sqlite_master;").fetchall()]

assert_in('Courses', tables, 'Checking the table has the correct name')

# Checking the column names
cur.execute("SELECT * FROM Courses")
columns = {c[0] for c in cur.description}

assert_equal(columns, {'CourseId','CourseNumber','CourseSubject','InstructorId','Section','Term','Year'})

con.close()
```

<!-- grade_id: 8bfa39, points: 1, grade: true, solution: false -->
```python
# Tests for creating the table CourseMeetings
import sqlite3
from nose.tools import assert_equal, assert_in

# Setting up sqlite3 for unit tests
con = sqlite3.connect(':memory:')
cur = con.cursor()

# Running all the code in act6create.sql
with open('act6create.sql') as file:
    cur.executescript(file.read())

# Checking that the table exists 
tables = [t[0] for t in cur.execute("SELECT tbl_name FROM sqlite_master;").fetchall()]

assert_in('CourseMeetings', tables, 'Checking the table has the correct name')

# Checking the column names
cur.execute("SELECT * FROM CourseMeetings")
columns = {c[0] for c in cur.description}

assert_equal(columns, {'CourseId', 'CourseMeetingId', 'Days', 'EndTime', 'Location', 'StartTime'})

con.close()
```

<!-- grade_id: 140893, points: 1, grade: true, solution: false -->
```python
# Tests for creating the table CourseStudent
import sqlite3
from nose.tools import assert_equal, assert_in

# Setting up sqlite3 for unit tests
con = sqlite3.connect(':memory:')
cur = con.cursor()

# Running all the code in act6create.sql
with open('act6create.sql') as file:
    cur.executescript(file.read())

# Checking that the table exists 
tables = [t[0] for t in cur.execute("SELECT tbl_name FROM sqlite_master;").fetchall()]

assert_in('CourseStudent', tables, 'Checking the table has the correct name')

# Checking the column names
cur.execute("SELECT * FROM CourseStudent")
columns = {c[0] for c in cur.description}

assert_equal(columns, {'CourseId', 'StudentId'})

con.close()
```

<!-- grade_id: 3140ca, points: 1, grade: true, solution: false -->
```python
# Tests for creating the table CourseDetails
import sqlite3
from nose.tools import assert_equal, assert_in

# Setting up sqlite3 for unit tests
con = sqlite3.connect(':memory:')
cur = con.cursor()

# Running all the code in act6create.sql
with open('act6create.sql') as file:
    cur.executescript(file.read())

# Checking that the table exists 
tables = [t[0] for t in cur.execute("SELECT tbl_name FROM sqlite_master;").fetchall()]

assert_in('CourseDetails', tables, 'Checking the table has the correct name')

# Checking the column names
cur.execute("SELECT * FROM CourseDetails")
columns = {c[0] for c in cur.description}

assert_equal(columns, {'CourseDescription', 'CourseNumber', 'CourseSubject', 'Credits', 'Title'})

con.close()
```

## Tests: Inserting Data

<!-- grade_id: b80e75, points: 1, grade: true, solution: false -->
```python
# Testing the sample data inserted into Courses
import sqlite3
from nose.tools import assert_in

# Setting up sqlite3 for unit tests
con = sqlite3.connect(':memory:')
cur = con.cursor()

# Running all the code in act6create.sql
with open('act6create.sql') as file:
    cur.executescript(file.read())
    
# Running all the code in act6insert.sql
with open('act6insert.sql') as file:
    cur.executescript(file.read())

# Checking for sample data 
rows = cur.execute("SELECT CourseId, CourseNumber, CourseSubject, InstructorId, Section, Term, Year FROM Courses;").fetchall()
con.close()

row = (1234, '270', 'CSC', 256, 1, 'Spring', '2024')
assert_in(row, rows, 'Checking one row exists')

row = (2468, '271L', 'PHY', 12, 2, 'Fall', '2024')
assert_in(row, rows, 'Checking another row exists')

# There are hidden tests which will *completely* check your solution.
```

<!-- grade_id: eadacc, points: 1, grade: true, solution: false -->
```python
# Testing the sample data inserted into CourseMeetings
import sqlite3
from nose.tools import assert_in

# Setting up sqlite3 for unit tests
con = sqlite3.connect(':memory:')
cur = con.cursor()

# Running all the code in act6create.sql
with open('act6create.sql') as file:
    cur.executescript(file.read())
    
# Running all the code in act6insert.sql
with open('act6insert.sql') as file:
    cur.executescript(file.read())

# Checking for sample data 
rows = cur.execute("SELECT CourseMeetingId, CourseId, Days, StartTime, EndTime, Location FROM CourseMeetings;").fetchall()
con.close()

row = (12, 1012, 'MWF', '13:00', '13:50', 'Boone 132')
assert_in(row, rows, 'Checking one row exists')

# There are hidden tests which will *completely* check your solution.
```

<!-- grade_id: ea09db, points: 1, grade: true, solution: false -->
```python
# Testing the sample data inserted into CourseStudent
import sqlite3
from nose.tools import assert_in

# Setting up sqlite3 for unit tests
con = sqlite3.connect(':memory:')
cur = con.cursor()

# Running all the code in act6create.sql
with open('act6create.sql') as file:
    cur.executescript(file.read())
    
# Running all the code in act6insert.sql
with open('act6insert.sql') as file:
    cur.executescript(file.read())

# Checking for sample data 
rows = cur.execute("SELECT StudentId, CourseId FROM CourseStudent;").fetchall()
con.close()

assert_in((201, 1111), rows, 'Checking one row exists')

# There are hidden tests which will *completely* check your solution.
```

<!-- grade_id: b895aa, points: 1, grade: true, solution: false -->
```python
# Tests for part 6
import sqlite3
from nose.tools import assert_in

# Setting up sqlite3 for unit tests
con = sqlite3.connect(':memory:')
cur = con.cursor()

# Running all the code in act6create.sql
with open('act6create.sql') as file:
    cur.executescript(file.read())
    
# Running all the code in act6insert.sql
with open('act6insert.sql') as file:
    cur.executescript(file.read())

# Checking for sample data 
rows = cur.execute("""
                   SELECT y.FirstName, y.LastName, c.Year
                   FROM Yotes AS y JOIN Courses AS c ON c.InstructorId = y.YoteId
                   WHERE c.CourseSubject = 'CSC'
                        AND c.CourseNumber = '270'
                        AND c.Section = 1
                        AND c.Term = 'Spring'
                        
                   """).fetchall()
con.close()

assert_in(('Robin', 'Cruz', '2021'), rows, 'Checking for Robin Cruz')
assert_in(('Jason', 'McClelland', '2020'), rows, 'Checking for Jason McClelland')
```

## Tests: Querying the database

<!-- grade_id: 589937, points: 2, grade: true, solution: false -->
```python
# Tests for part 7
import sqlite3
from nose.tools import assert_equal
from csc270tools import get_script

# Setting up sqlite3 for unit tests
con = sqlite3.connect(':memory:')
cur = con.cursor()

# Running all the code in act6create.sql
with open('act6create.sql') as file:
    cur.executescript(file.read())
    
# Running all the code in act6insert.sql
with open('act6insert.sql') as file:
    cur.executescript(file.read())
    
# executing the SQL query in act6query.sql for this part
rows = cur.execute(get_script('act6query.sql', 0)).fetchall()

con.close()

assert_equal(rows, [('Jonny', 'Comes'), ('Robin', 'Cruz'), ('Jason', 'McClelland')], 'Checking the rows')
```

### Grading the last part

The following cell runs your solution for part 8. If your solution is correct, the cell should run without an error **and** your course schedule will be printed at the bottom of this notebook. I will manually check your course schedule by looking at the printout at the bottom of this notebook. 

<!-- grade_id: a7c7ce, points: 1, grade: true, solution: false -->
```python
# Tests for part 8
import sqlite3
from nose.tools import assert_equal
from csc270tools import get_script

# Setting up sqlite3 for unit tests
con = sqlite3.connect(':memory:')
cur = con.cursor()

# Running all the code in act6create.sql
with open('act6create.sql') as file:
    cur.executescript(file.read())
    
# Running all the code in act6insert.sql
with open('act6insert.sql') as file:
    cur.executescript(file.read())
    
# executing the SQL query in act6query.sql for this part
rows = cur.execute(get_script('act6query.sql', 1)).fetchall()
header = [c[0] for c in cur.description]
con.close()

assert_equal(header, ['course', 'course_title', 'instructor', 'days', 'start_time', 'end_time', 'room'], 'Checking the column names')

# Print the output of your query for manual grading
a,b,c,d,e,f,g = header
print(f"\033[1m{a:14}{b:35}{c:20}{d:6}{e:12}{f:12}{g}\033[0m")
for a,b,c,d,e,f,g in rows:
    print(f"{a:14}{b:35}{c:20}{d:6}{e:12}{f:12}{g}")
```

Output:
```
[1mcourse        course_title                       instructor          days  start_time  end_time    room[0m
CSC-270-01    Applied Databases                  Comes               TH    09:25       10:40       CML 105
CSC-340-01    Design and Analysis of Algorithms  Danielson           TT    8:00        09:15       CML 106
MAT-361-01    Linear Algebra                     Rosoff              MWF   11:00       11:50       Boone 122
ARH-205-01    Renaissance to Modern              Jones               MWF   15:00       17:50       Langroise 201
ART-375-02    Digital Video                      Seyben              MWF   14:00       14:50       Langroise 202
```