# Tests for Activity 1


## Test: Hello, World!

The following cell contains code that tests the contents of your `hello.sql` file. Run the following cell and see what happens. If the cell runs without errors, you have successfully completed the first SQL exercise for this course (Good job!). In other words, no news is good news.

If you get an error when you run the following cell, you should read the error message carefully (start with the bottom of the error message) and try to fix your mistake. If you can't figure out what's wrong, ask for help!

<!-- grade_id: 0c9495, points: 1, grade: true, solution: false -->
```python
import sqlite3
from nose.tools import assert_equal

# Setting up sqlite3 for unit tests
con = sqlite3.connect(':memory:')
cur = con.cursor()

# Running the code in hello.sql
with open('hello.sql') as file:
    query = file.read()
    table = cur.execute(query).fetchall()
    
# Checking the output of the query
assert_equal(table[0][0], 'Hello, World!')
```

## Test: Chinook

The following cell contains code that **partially** tests the query in your `employees.sql` file. If you see errors when you run the cell below, you definitely have a mistake to fix, and you should take the appropriate steps to fix your code.

In addition to the tests you can see, there are hidden tests that will fully test your code for grading purposes. So even if you don't see any errors, you'll still want to double check your work to **make sure your code properly satisfies the directions in the handout**.

<!-- grade_id: 5b7c7d, points: 1, grade: true, solution: false -->
```python
import sqlite3
from nose.tools import assert_equal

# Setting up sqlite3 for unit tests
con = sqlite3.connect('Chinook.db')
cur = con.cursor()

# Running the code in employees.sql
with open('employees.sql') as file:
    query = file.read()
    table = cur.execute(query).fetchall()
    
# Partially checking the output of the query
first_row = table[0]
assert_equal(first_row, ('Andrew', 'Adams', 'General Manager'), 'Checking the first row of the output')

# There are hidden tests which will *completely* check your solution.
```

## Test: Course Schedule

<!-- grade_id: 93679d, points: 1, grade: true, solution: false -->
```python
import sqlite3
from nose.tools import assert_equal

# Setting up sqlite3 for unit tests
con = sqlite3.connect(':memory:')
cur = con.cursor()

# Running the code in schedule.sql
with open('schedule.sql') as file:
    contents = file.read()
    for query in contents.split(';'):
        cur.execute(query)
    
# Checking that a table named course_schedule exists
tables = cur.execute('''
    SELECT name FROM sqlite_master WHERE type = 'table';
    ''').fetchall()
assert_equal(tables[0][0], 'course_schedule')
```

Output:
```
IndexError                                Traceback (most recent call last)
/tmp/ipykernel_1344/2777610781.py in <cell line: 18>()
     16     SELECT name FROM sqlite_master WHERE type = 'table';
     17     ''').fetchall()
---> 18 assert_equal(tables[0][0], 'course_schedule')

IndexError: list index out of range
```

```python
from IPython.display import Markdown, display

display(Markdown("schedule.md"))
```

Output:
```
<IPython.core.display.Markdown object>
```