# Tests for Activity 2 

This notebook tests some of the code you have written for the activity. Usually, the tests you will see only partially check your code. In addition to the visible tests, there are hidden tests that fully check your code. Make sure your queries give you precisely what the directions ask for!

## Test: Practice with Chinook

<!-- grade_id: 712baa, points: 1, grade: true, solution: false -->
```python
# Test for Part 2
from nose.tools import assert_equal
from csc270tools import get_table

# Running the code in prac_chinook.sql to get the table for this part
table = get_table('prac_chinook.sql', 'Chinook.db', 0)

assert_equal(len(table), 10, 'Checking the number of rows')
assert_equal(table[0], (1, 'For Those About To Rock We Salute You', 1), 'Checking the first row')

```

<!-- grade_id: 151b3f, points: 1, grade: true, solution: false -->
```python
# Test for Part 3
from nose.tools import assert_equal
from csc270tools import get_table

# Running the code in prac_chinook.sql to get the table for this part
table = get_table('prac_chinook.sql', 'Chinook.db', 1)

assert_equal(table, [('Zeca Pagodinho',), ("Youssou N'Dour",)])
```

<!-- grade_id: 9c137c, points: 1, grade: true, solution: false -->
```python
# Test for 4(a)
from nose.tools import assert_equal
from csc270tools import get_table

# Running the code in prac_chinook.sql to get the table for this part
table = get_table('prac_chinook.sql', 'Chinook.db', 2)

assert_equal(table[0], ('Camille', 'Bernard', '4, Rue Milton', '75009'), 'Checking the first row')

```

<!-- grade_id: b5da0e, points: 1, grade: true, solution: false -->
```python
# Test for 4(b)
from nose.tools import assert_equal, assert_less
from csc270tools import get_table

# Running the code in prac_chinook.sql to get the table for this part
table = get_table('prac_chinook.sql', 'Chinook.db', 3)

assert_less(len(table), 8, 'Checking that there are fewer than 8 rows')
assert_equal(table[0], ('Margaret', 'Park', 'Sales Support Agent', '1947-09-19 00:00:00'), 'Checking the first row')

```

<!-- grade_id: 4707ab, points: 1, grade: true, solution: false -->
```python
# Test for 4(c)
from nose.tools import assert_equal
from csc270tools import get_table

# Running the code in prac_chinook.sql to get the table for this part
table = get_table('prac_chinook.sql', 'Chinook.db', 4)

assert_equal(table[0], ('Laura', 'Callahan', 'IT Staff', '2004-03-04 00:00:00'), 'Checking the first row')
assert_equal(table[-1], ('Margaret', 'Park', 'Sales Support Agent', '2003-05-03 00:00:00'), 'Checking the last row')

```

<!-- grade_id: c9ff27, points: 1, grade: true, solution: false -->
```python
# Test for 4(d)
from nose.tools import assert_equal
from csc270tools import get_table

# Running the code in prac_chinook.sql to get the table for this part
table = get_table('prac_chinook.sql', 'Chinook.db', 5)

assert_equal(len(table), 1, 'Checking the table has 1 row')
assert_equal(len(table[0]), 1, 'Checking the table has 1 column')

```

<!-- grade_id: d2c8d7, points: 1, grade: true, solution: false -->
```python
# Test for 4(e)
from nose.tools import assert_equal
from csc270tools import get_table

# Running the code in prac_chinook.sql to get the table for this part
table = get_table('prac_chinook.sql', 'Chinook.db', 6)

assert_equal(len(table), 4, 'Checking the number of rows')
assert_equal(table[0], ('Diver Down',), 'Checking the first row')

```

## Tests: Practice with Northwind

<!-- grade_id: fdfef9, points: 1, grade: true, solution: false -->
```python
# Test for part 1
from nose.tools import assert_equal, assert_less
from csc270tools import get_table

# Running the code in prac_northwind.sql to get the table for this part
table = get_table('prac_northwind.sql', 'Northwind.db', 0)

assert_equal(table[0], ('Sales Representative',), 'Checking the first row')
assert_less(len(table), 9, 'Checking that there are fewer than 9 rows')

```

<!-- grade_id: 3ce006, points: 1, grade: true, solution: false -->
```python
# Test for part 2
from nose.tools import assert_equal
from csc270tools import get_table

# Running the code in prac_northwind.sql to get the table for this part
table = get_table('prac_northwind.sql', 'Northwind.db', 1)

assert_equal(table[0], ('Laura', 'Callahan', 'Inside Sales Coordinator'), 'Checking the first row')

```

<!-- grade_id: d5900c, points: 1, grade: true, solution: false -->
```python
# Test for part 3
from nose.tools import assert_equal, assert_less
from csc270tools import get_table

# Running the code in prac_northwind.sql to get the table for this part
table = get_table('prac_northwind.sql', 'Northwind.db', 2)

assert_equal(len(table[0]), 1, 'Checking that there is a single column')
assert_equal(table[0][0], 'Argentina', 'Checking the first row')
assert_equal(table[-1][0], 'Venezuela', 'Checking the last row')
assert_less(len(table), 100, 'Checking that there are fewer than 100 rows')

```

<!-- grade_id: 187076, points: 1, grade: true, solution: false -->
```python
# Test for part 4
from nose.tools import assert_equal
from csc270tools import get_table

# Running the code in prac_northwind.sql to get the table for this part
table = get_table('prac_northwind.sql', 'Northwind.db', 3)

assert_equal(table[0], (10264, '2016-07-24', '2016-08-21', '2016-08-23'), 'Checking the first row')
assert_equal(len(table), 37, 'Checking the number of rows')


```

<!-- grade_id: 8d367d, points: 1, grade: true, solution: false -->
```python
# Test for part 5
from nose.tools import assert_equal, assert_less
from csc270tools import get_table

# Running the code in prac_northwind.sql to get the table for this part
table = get_table('prac_northwind.sql', 'Northwind.db', 4)

assert_equal(table[0], ("Chef Anton's Gumbo Mix", 0), 'Checking the first row')
assert_equal(table[1], ('Mishi Kobe Niku', 29), 'Checking the second row')
assert_equal(len(table), 8, 'Checking the number of rows')


```

<!-- grade_id: 74295c, points: 1, grade: true, solution: false -->
```python
# Test for part 6
from nose.tools import assert_equal
from csc270tools import get_table

# Running the code in prac_northwind.sql to get the table for this part
table = get_table('prac_northwind.sql', 'Northwind.db', 5)

assert_equal(table, [(10248, 11, 14, 12, 0.0),
                     (10248, 42, 9.8, 10, 0.0),
                     (10248, 72, 34.8, 5, 0.0),
                     (10249, 14, 18.6, 9, 0.0)], 'Checking the whole table')


```

<!-- grade_id: fd1060, points: 1, grade: true, solution: false -->
```python
# Test for part 7
from nose.tools import assert_equal
from csc270tools import get_table

# Running the code in prac_northwind.sql to get the table for this part
table = get_table('prac_northwind.sql', 'Northwind.db', 6)

assert_equal(table[0], (10850, 4, 2.5), 'Checking the first row')
assert_equal(table[-1], (10382, 60, 2), 'Checking the last row')
assert_equal(len(table), 32, 'Checking the number of rows')


```

## Tests: Practice with Netflix

<!-- grade_id: 4b9cb7, points: 1, grade: true, solution: false -->
```python
# Test for part 4
from nose.tools import assert_equal
from csc270tools import get_table

# Running the code in prac_netflix.sql to get the table for this part
table = get_table('prac_netflix.sql', 'Netflix.db', 1)

assert_equal(table[0], ('Django Unchained', 'R', '165 min', '2012'), 'Checking the first row')


```

<!-- grade_id: cfdcfb, points: 1, grade: true, solution: false -->
```python
# Test for part 5
from nose.tools import assert_equal
from csc270tools import get_table

# Running the code in prac_netflix.sql to get the table for this part
table = get_table('prac_netflix.sql', 'Netflix.db', 2)

assert_equal(table[0], ('An American Tail: The Treasures of Manhattan Island', 'Movie', 'G', '79 min'), 'Checking the first row')
assert_equal(table[1], ('Rugrats in Paris: The Movie', 'Movie', 'G', '79 min'), 'Checking the second row')
assert_equal(table[-1], ('Power Rangers Lightspeed Rescue', 'TV Show', 'TV-Y7', '1 Season'), 'Checking the last row')


```

<!-- grade_id: 0cdc8b, points: 1, grade: true, solution: false -->
```python
# Test for part 6a
from nose.tools import assert_equal
from csc270tools import get_table

# Running the code in prac_netflix.sql to get the table for this part
table = get_table('prac_netflix.sql', 'Netflix.db', 3)

assert_equal(len(table), 25, 'Checking number of rows')
assert_equal(table[0], ('Blood Brothers: Malcolm X & Muhammad Ali', 'September 9, 2021'), 'Checking the first row')


```

```python

```