# Tests for Activity 3

## Tests: Northwind

<!-- grade_id: ce5c0f, points: 2, grade: true, solution: false -->
```python
# Test for part 1
from nose.tools import assert_equal
from csc270tools import get_table

# Running the code in act3Northwind.sql to get the table for this part
table = get_table('act3Northwind.sql', 'Northwind.db', 0)

assert_equal(table[0], (11008, '2018-04-08', '2018-05-06'), 'Checking the first row')

```

<!-- grade_id: bacf02, points: 2, grade: true, solution: false -->
```python
# Test for part 2
from nose.tools import assert_equal
from csc270tools import get_table

# Running the code in act3Northwind.sql to get the table for this part
table = get_table('act3Northwind.sql', 'Northwind.db', 1)

assert_equal(table[0], ('BERGS', '2016-08-16', 'Luleå'), 'Checking the first row')
assert_equal(table[-1], ('SIMOB', None, 'Kobenhavn'), 'Checking the last row')

```

<!-- grade_id: aae28a, points: 2, grade: true, solution: false -->
```python
# Test for part 3
from nose.tools import assert_equal
from csc270tools import get_table

# Running the code in act3Northwind.sql to get the table for this part
table = get_table('act3Northwind.sql', 'Northwind.db', 2)

assert_equal(table[0], ('MORGK', '0342-023176', None, 'address'), 'Checking the first row')
assert_equal(table[-1], ('VALON', None, None, 'both'), 'Checking the last row')

```

## Tests: Netflix

<!-- grade_id: 602cd6, points: 1, grade: true, solution: false -->
```python
# Test for part 1
from nose.tools import assert_equal
from csc270tools import get_table

# Running the code in act3Netflix.sql to get the table for this part
table = get_table('act3Netflix.sql', 'Netflix.db', 0)

assert_equal(table[0], ('The Beguiled', 'Sofia Coppola'), 'Checking the first row')
assert_equal(len(table), 5, 'Checking the number of rows')

```

<!-- grade_id: df5f99, points: 1, grade: true, solution: false -->
```python
# Test for part 2
from nose.tools import assert_equal
from csc270tools import get_table

# Running the code in act3Netflix.sql to get the table for this part
table = get_table('act3Netflix.sql', 'Netflix.db', 1)

assert_equal(table[0], ('Bygones Be Bygones', '2017'), 'Checking the first row')
assert_equal(len(table), 13, 'Checking the number of rows')

```

<!-- grade_id: 560aa7, points: 1, grade: true, solution: false -->
```python
# Test for part 3
from nose.tools import assert_equal
from csc270tools import get_table

# Running the code in act3Netflix.sql to get the table for this part
table = get_table('act3Netflix.sql', 'Netflix.db', 2)

assert_equal(table[0], ('The Walking Dead', '10 Seasons'), 'Checking the first row')
assert_equal(len(table), 17, 'Checking the number of rows')

```

<!-- grade_id: be256c, points: 1, grade: true, solution: false -->
```python
# Test for part 4
from nose.tools import assert_equal
from csc270tools import get_table

# Running the code in act3Netflix.sql to get the table for this part
table = get_table('act3Netflix.sql', 'Netflix.db', 3)

assert_equal(table[0], ('Bob Ross: Happy Accidents, Betrayal & Greed',), 'Checking the first row')
assert_equal(len(table), 17, 'Checking the number of rows')

```

<!-- grade_id: 8a994a, points: 2, grade: true, solution: false -->
```python
# Test for part 5
from nose.tools import assert_equal
from csc270tools import get_table

# Running the code in act3Netflix.sql to get the table for this part
table = get_table('act3Netflix.sql', 'Netflix.db', 4)

assert_equal(table[0], ('Between Worlds', 'Thrillers'), 'Checking the first row')
assert_equal(len(table), 12, 'Checking the number of rows')

```

<!-- grade_id: a46727, points: 2, grade: true, solution: false -->
```python
# Test for part 6
from nose.tools import assert_equal
from csc270tools import get_table

# Running the code in act3Netflix.sql to get the table for this part
table = get_table('act3Netflix.sql', 'Netflix.db', 5)

assert_equal(table[0], ('Scream 2', '1997', 'Horror Movies'), 'Checking the first row')
assert_equal(len(table), 9, 'Checking the number of rows')

```

<!-- grade_id: 315dc1, points: 2, grade: true, solution: false -->
```python
# Test for part 7
from nose.tools import assert_equal
from csc270tools import get_table

# Running the code in act3Netflix.sql to get the table for this part
table = get_table('act3Netflix.sql', 'Netflix.db', 6)

assert_equal(table[0], ('Sankofa', 0), 'Checking the first row')
assert_equal(table[5], ("Beethoven's 2nd", 1), 'Checking another row')
assert_equal(len(table), 28, 'Checking the number of rows')

```

<!-- grade_id: d290ac, points: 2, grade: true, solution: false -->
```python
# Test for part 8
from nose.tools import assert_equal
from csc270tools import get_table

# Running the code in act3Netflix.sql to get the table for this part
table = get_table('act3Netflix.sql', 'Netflix.db', 7)

assert_equal(table[0], ('Hard Lessons (1986)',), 'Checking the first row')
assert_equal(len(table), 11, 'Checking the number of rows')

```